// Numbas version: finer_feedback_settings
{"name": "R\u00f8tter", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": true, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Kvadratr\u00f8tter", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Kvadratr\u00f8tter", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om kvadratr&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(4..9)^2", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "100*random(5..8)^2", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "clist[index1]", "description": "", "templateType": "anything", "can_override": false}, "clist": {"name": "clist", "group": "Ungrouped variables", "definition": "[98, 75, 80, 300, 128, 45]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d1sqrt": {"name": "d1sqrt", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "d2sqrt": {"name": "d2sqrt", "group": "Ungrouped variables", "definition": "random(3..9 except d1sqrt)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "e2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Ungrouped variables", "definition": "[21, 26, 60, 27, 33, 30]", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "e1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Ungrouped variables", "definition": "[56, 39, 48, 45, 77,54]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b1", "b", "c", "clist", "index1", "d1", "d2", "d1sqrt", "d2sqrt", "e2", "e1list", "e1", "e2list"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\var{a}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Du m&aring; finne det positive tallet $x$ slik at $x&sup2;=\\var{a}$.</p>"}], "minValue": "sqrt(a)", "maxValue": "sqrt(a)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\var{b}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Du m&aring; finne det positive tallet $x$ slik at $x&sup2;=\\var{b}$.</p>"}], "minValue": "b^{1/2}", "maxValue": "b^{1/2}", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\var{c}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Tallet $\\var{c}$ kan skrives som et produkt av $x&sup2; \\cdot y,$ der $x$ og $y$ er to positive tall. Hva er $x$ og $y$?</p>\n<p>Husk at $\\sqrt{x&sup2;\\cdot y} = x\\cdot \\sqrt{y}$.</p>\n<p>For eksempel:</p>\n<p>$\\sqrt{245} = \\sqrt{49 \\cdot 5} = \\sqrt{7&sup2; \\cdot 5} = 7\\sqrt{5}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$7\\sqrt{2}$", "$5\\sqrt{3}$", "$4\\sqrt{5}$", "$10\\sqrt{3}$", "$8\\sqrt{2}$", "$3\\sqrt{5}$", "$9\\sqrt{2}$", "$6\\sqrt{3}$", "$2\\sqrt{5}$", "$2\\sqrt{7}$", "$3\\sqrt{10}$", "$2\\sqrt{8}$", "$3\\sqrt{6}$", "$5\\sqrt{2}$"], "matrix": ["if(index1=0,1,0)", "if(index1=1,1,0)", "if(index1=2,1,0)", "if(index1=3,1,0)", "if(index1=4,1,0)", "if(index1=5,1,0)", 0, 0, 0, 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", "", "", "", "", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Kvadratr\u00f8tter niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om kvadratr&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(4..9)^2", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1^3", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "clist[index1]", "description": "", "templateType": "anything", "can_override": false}, "clist": {"name": "clist", "group": "Ungrouped variables", "definition": "[98, 75, 80, 300, 128, 45]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d1sqrt": {"name": "d1sqrt", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "d2sqrt": {"name": "d2sqrt", "group": "Ungrouped variables", "definition": "random(3..9 except d1sqrt)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "e2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Ungrouped variables", "definition": "[21, 26, 60, 27, 33, 30]", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "e1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Ungrouped variables", "definition": "[56, 39, 48, 45, 77,54]", "description": "", "templateType": "anything", "can_override": false}, "a1": {"name": "a1", "group": "Ungrouped variables", "definition": "a2*ak", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "ak": {"name": "ak", "group": "Ungrouped variables", "definition": "random(4..9)^2", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b1", "b", "c", "clist", "index1", "d1", "d2", "d1sqrt", "d2sqrt", "e2", "e1list", "e1", "e2list", "a1", "a2", "ak"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\frac{\\var{d1}}{\\var{d2}}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Husk at $\\sqrt{\\frac{\\var{d1}}{\\var{d2}}} = \\frac{\\sqrt{\\var{d1}}}{\\sqrt{\\var{d2}}} $</p>"}], "minValue": "d1sqrt/d2sqrt", "maxValue": "d1sqrt/d2sqrt", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\frac{\\var{a1}}{\\var{a2}}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Pr&oslash;v f&oslash;rst &aring; forkorte br&oslash;ken. F&aring;r du et kvadrat hvis du deler $\\var{a1}$ med&nbsp; $\\var{a2}$?</p>"}], "minValue": "ak^{1/2}", "maxValue": "ak^{1/2}", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt{\\frac{\\var{e1}}{\\var{e2}}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Pr&oslash;v &aring; begynne med &aring; forkorte br&oslash;ken f&oslash;rst. Hva f&aring;r du hvis du forkorter $\\frac{\\var{e1}}{\\var{e2}}?$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\frac{\\sqrt{3}}{\\sqrt{8}}$", "$\\frac{\\sqrt{2}}{\\sqrt{3}}$", "$\\frac{\\sqrt{5}}{2}$", "$\\frac{\\sqrt{3}}{\\sqrt{5}}$", "$\\frac{\\sqrt{3}}{\\sqrt{7}}$", "$\\frac{\\sqrt{5}}{3}$", "$\\frac{\\sqrt{3}}{2}$", "$\\frac{\\sqrt{2}}{\\sqrt{9}}$", "$\\frac{\\sqrt{5}}{\\sqrt{7}}$"], "matrix": ["if(index1=0, 1,0)", "if(index1=1, 1,0)", "if(index1=2, 1,0)", "if(index1=3, 1,0)", "if(index1=4, 1,0)", "if(index1=5, 1,0)", 0, 0, 0], "distractors": ["", "", "", "", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Tredje r\u00f8tter", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Tredje r\u00f8tter", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om tredje r&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "-arad^3", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1^3", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "clist[index1]", "description": "", "templateType": "anything", "can_override": false}, "clist": {"name": "clist", "group": "Ungrouped variables", "definition": "[98, 75, 80, 300, 128, 45]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2sqrt^2", "description": "", "templateType": "anything", "can_override": false}, "d1sqrt": {"name": "d1sqrt", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "d2sqrt": {"name": "d2sqrt", "group": "Ungrouped variables", "definition": "random(3..9 except d1sqrt)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "e2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Ungrouped variables", "definition": "[21, 26, 60, 27, 33, 30]", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "e1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Ungrouped variables", "definition": "[56, 39, 48, 45, 77,54]", "description": "", "templateType": "anything", "can_override": false}, "arad": {"name": "arad", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b1", "b", "c", "clist", "index1", "d1", "d2", "d1sqrt", "d2sqrt", "e2", "e1list", "e1", "e2list", "arad"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt[3]{\\var{b}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Du m&aring; finne tallet $x$ slik at $x^3=\\var{b}$.</p>"}], "minValue": "b1", "maxValue": "b1", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt[3]{\\var{a}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Du m&aring; finne tallet $x$ slik at $x^3=\\var{a}$. Er $x$ positivt eller negativt?</p>"}], "minValue": "-arad", "maxValue": "-arad", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Tredje r\u00f8tter niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om tredje r&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "a3*a2", "description": "", "templateType": "anything", "can_override": false}, "a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(2..3)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..5 except a1)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "a1^3", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "c11^3", "description": "", "templateType": "anything", "can_override": false}, "c11": {"name": "c11", "group": "Ungrouped variables", "definition": "-random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "c22": {"name": "c22", "group": "Ungrouped variables", "definition": "random(2..5 except -c11)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "c22^3", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "a1", "a2", "a3", "c1", "c11", "c22", "c2"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt[3]{\\var{a}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Tallet $\\var{a}$ kan skrives som et produkt av $x^3 \\cdot y,$ der $x$ og $y$ er to tall. Hva er $x$ og $y$?</p>\n<p>Husk at $\\sqrt[3]{x^3\\cdot y} = x\\cdot \\sqrt[3]{y}$.</p>\n<p>For eksempel:</p>\n<p>$\\sqrt[3]{128} = \\sqrt[3]{64 \\cdot 2} = \\sqrt[3]{4^3 \\cdot 2} = 4\\sqrt[3]{2}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{a1*a2}$", "$\\var{a1}\\cdot\\sqrt[3]{\\var{a2}}$", "$\\var{a2}\\cdot\\sqrt[3]{\\var{a1}}$", "$8\\sqrt[3]{\\var{a1}}$", "$8\\sqrt[3]{\\var{a2}}$", "$10\\sqrt[3]{\\var{a2}}$", "$10\\sqrt[3]{\\var{a1}}$"], "matrix": ["0", "1", 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", "", ""]}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\sqrt[3]{\\frac{\\var{c1}}{\\var{c2}}}$. Skriv svaret som en br&oslash;k.</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Husk at $\\sqrt[3]{\\frac{\\var{c1}}{\\var{c2}}} = \\frac{\\sqrt[3]{\\var{c1}}}{\\sqrt[3]{\\var{c2}}} $</p>\n<p>Pass p&aring; fortegnet av $\\sqrt[3]{\\var{c1}}$!</p>"}], "minValue": "c11/c22", "maxValue": "c11/c22", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "N-te r\u00f8tter", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", ""], "variable_overrides": [[], [], []], "questions": [{"name": "N-te r\u00f8tter", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgaver om r&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "listoflist[index1] ", "description": "", "templateType": "anything", "can_override": false}, "esp": {"name": "esp", "group": "Part a", "definition": "espl[index1]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "espl": {"name": "espl", "group": "Part a", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "list2": {"name": "list2", "group": "Part a", "definition": "random(5..11)^2", "description": "", "templateType": "anything", "can_override": false}, "listoflist": {"name": "listoflist", "group": "Part a", "definition": "[list2,list3]", "description": "", "templateType": "anything", "can_override": false}, "list3": {"name": "list3", "group": "Part a", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "blist": {"name": "blist", "group": "Part b", "definition": "[3,2,2,4,4,3]", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part b", "definition": "blist[index2]", "description": "", "templateType": "anything", "can_override": false}, "rtlist": {"name": "rtlist", "group": "Part b", "definition": "[2,3,4,2,3,4]", "description": "", "templateType": "anything", "can_override": false}, "rt": {"name": "rt", "group": "Part b", "definition": "rtlist[index2]", "description": "", "templateType": "anything", "can_override": false}, "index2": {"name": "index2", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Part c", "definition": "c1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Part c", "definition": "c2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "listc2": {"name": "listc2", "group": "Part c", "definition": "random(2..5)^2", "description": "", "templateType": "anything", "can_override": false}, "listc3": {"name": "listc3", "group": "Part c", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "listoflistc": {"name": "listoflistc", "group": "Part c", "definition": "[listc3, listc2]", "description": "", "templateType": "anything", "can_override": false}, "c1list": {"name": "c1list", "group": "Part c", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part c", "definition": "listoflistc[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2list": {"name": "c2list", "group": "Part c", "definition": "[3,2]", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Part e", "definition": "random(2..9 except d3)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d4": {"name": "d4", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d5": {"name": "d5", "group": "Part e", "definition": "random(2..9 except d4)", "description": "", "templateType": "anything", "can_override": false}, "d6": {"name": "d6", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "sold": {"name": "sold", "group": "Part e", "definition": "d1*d2/d3 + d4/d5*d6", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Part d", "definition": "e1list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Part d", "definition": "e2list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Part d", "definition": "e3list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Part d", "definition": "[3,2,4,3,3,4]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Part d", "definition": "[25,125,8,4,9,27]", "description": "", "templateType": "anything", "can_override": false}, "e3list": {"name": "e3list", "group": "Part d", "definition": "[2,3,2,4,2,2]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["index1", "index2"], "variable_groups": [{"name": "Part a", "variables": ["a", "esp", "espl", "list2", "list3", "listoflist"]}, {"name": "Part b", "variables": ["b", "blist", "rt", "rtlist"]}, {"name": "Part c", "variables": ["c1", "c2", "listc2", "listc3", "listoflistc", "c1list", "c2list", "c"]}, {"name": "Part e", "variables": ["d1", "d2", "d3", "d4", "d5", "d6", "sold"]}, {"name": "Part d", "variables": ["e1", "e2", "e3", "e1list", "e2list", "e3list"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\var{a}^{\\frac{1}{\\var{esp}}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Husk at $\\var{a}^{\\frac{1}{\\var{esp}}} = \\sqrt[\\var{esp}]{\\var{a}}$</p>"}], "minValue": "a^(1/esp)", "maxValue": "a^(1/esp)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Beregn $\\var{c}^{\\var{c1}/\\var{c2}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Husk at $\\var{c}^{\\var{c1}/\\var{c2}} = \\sqrt[\\var{c2}]{\\var{c}^{\\var{c1}}} = \\Big(\\sqrt[\\var{c2}]{\\var{c}}\\Big)^{\\var{c1}}$</p>"}], "minValue": "c^{c1/c2}", "maxValue": "c^{c1/c2}", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "precisionType": "dp", "precision": 0, "precisionPartialCredit": 0, "precisionMessage": "You have not given your answer to the correct precision.", "strictPrecision": true, "showPrecisionHint": true, "notationStyles": ["plain", "plain-eu"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "N-te r\u00f8tter niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgaver om r&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "listoflist[index1] ", "description": "", "templateType": "anything", "can_override": false}, "esp": {"name": "esp", "group": "Part a", "definition": "espl[index1]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "espl": {"name": "espl", "group": "Part a", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "list2": {"name": "list2", "group": "Part a", "definition": "random(5..11)^2", "description": "", "templateType": "anything", "can_override": false}, "listoflist": {"name": "listoflist", "group": "Part a", "definition": "[list2,list3]", "description": "", "templateType": "anything", "can_override": false}, "list3": {"name": "list3", "group": "Part a", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "blist": {"name": "blist", "group": "Part b", "definition": "[3,2,2,4,4,3]", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part b", "definition": "blist[index2]", "description": "", "templateType": "anything", "can_override": false}, "rtlist": {"name": "rtlist", "group": "Part b", "definition": "[2,3,4,2,3,4]", "description": "", "templateType": "anything", "can_override": false}, "rt": {"name": "rt", "group": "Part b", "definition": "rtlist[index2]", "description": "", "templateType": "anything", "can_override": false}, "index2": {"name": "index2", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Part c", "definition": "c1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Part c", "definition": "c2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "listc2": {"name": "listc2", "group": "Part c", "definition": "random(2..5)^2", "description": "", "templateType": "anything", "can_override": false}, "listc3": {"name": "listc3", "group": "Part c", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "listoflistc": {"name": "listoflistc", "group": "Part c", "definition": "[listc3, listc2]", "description": "", "templateType": "anything", "can_override": false}, "c1list": {"name": "c1list", "group": "Part c", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part c", "definition": "listoflistc[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2list": {"name": "c2list", "group": "Part c", "definition": "[3,2]", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Part e", "definition": "random(2..9 except d3)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d4": {"name": "d4", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d5": {"name": "d5", "group": "Part e", "definition": "random(2..9 except d4)", "description": "", "templateType": "anything", "can_override": false}, "d6": {"name": "d6", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "sold": {"name": "sold", "group": "Part e", "definition": "d1*d2/d3 + d4/d5*d6", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Part d", "definition": "e1list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Part d", "definition": "e2list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Part d", "definition": "e3list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Part d", "definition": "[3,2,4,3,3,4]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Part d", "definition": "[25,125,8,4,9,27]", "description": "", "templateType": "anything", "can_override": false}, "e3list": {"name": "e3list", "group": "Part d", "definition": "[2,3,2,4,2,2]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["index1", "index2"], "variable_groups": [{"name": "Part a", "variables": ["a", "esp", "espl", "list2", "list3", "listoflist"]}, {"name": "Part b", "variables": ["b", "blist", "rt", "rtlist"]}, {"name": "Part c", "variables": ["c1", "c2", "listc2", "listc3", "listoflistc", "c1list", "c2list", "c"]}, {"name": "Part e", "variables": ["d1", "d2", "d3", "d4", "d5", "d6", "sold"]}, {"name": "Part d", "variables": ["e1", "e2", "e3", "e1list", "e2list", "e3list"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Omskriv $\\big(\\sqrt[\\var{rt}]{2}\\big)^{\\var{b}}$</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Husk at $\\sqrt[\\var{rt}]{2} = 2^{\\var{rt}}$.</p>\n<p>Husk ogs&aring; regnereglene for potenser:</p>\n<p>$(x^a)^b = x^{a\\cdot b}$. Hva skjer n&aring;r $a$ er en br&oslash;k?</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$2^{3/2}$", "$2^{2/3}$", "$2^{1/2}$", "$2^2$", "$2^{4/3}$", "$2^{3/4}$"], "matrix": ["if(index2=0,1,0)", "if(index2=1,1,0)", "if(index2=2,1,0)", "if(index2=3,1,0)", "if(index2=4,1,0)", "if(index2=5,1,0)"], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Tall $\\big( \\sqrt[\\var{e1}]{\\var{e2}} \\big)^{\\var{e3}}$ kan omskrives til:</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Pass p&aring; at $\\var{e2}$ kan omskrives til $x^y$, for noe tall $x$ og $y$.</p>\n<p>Etterp&aring;, husk &aring; bruke potensregler $(a^b)^c = a^{b\\cdot c}$</p>\n<p>For eksempel:</p>\n<p>$(\\sqrt[5]{16})^2 = (\\sqrt[5]{2^4})^2 = ((2^4)^{1/5})^2=2^{4 \\cdot \\frac{1}{5}\\cdot 2} = 2^{8/5}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$5^{4/3}$", "$5^{9/2}$", "$2^{3/2}$", "$2^{8/3}$", "$3^{4/3}$", "$3^{3/2}$", "$5^{2/3}$", "$2^{4/3}$", "$3^{4/2}$"], "matrix": ["if(index2=0,1,0)", "if(index2=1,1,0)", "if(index2=2,1,0)", "if(index2=3,1,0)", "if(index2=4,1,0)", "if(index2=5,1,0)", 0, 0, 0], "distractors": ["", "", "", "", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "N-te r\u00f8tter niv\u00e5 3", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgaver om r&oslash;tter.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "listoflist[index1] ", "description": "", "templateType": "anything", "can_override": false}, "esp": {"name": "esp", "group": "Part a", "definition": "espl[index1]", "description": "", "templateType": "anything", "can_override": false}, "index1": {"name": "index1", "group": "Ungrouped variables", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "espl": {"name": "espl", "group": "Part a", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "list2": {"name": "list2", "group": "Part a", "definition": "random(5..11)^2", "description": "", "templateType": "anything", "can_override": false}, "listoflist": {"name": "listoflist", "group": "Part a", "definition": "[list2,list3]", "description": "", "templateType": "anything", "can_override": false}, "list3": {"name": "list3", "group": "Part a", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "blist": {"name": "blist", "group": "Part b", "definition": "[3,2,2,4,4,3]", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part b", "definition": "blist[index2]", "description": "", "templateType": "anything", "can_override": false}, "rtlist": {"name": "rtlist", "group": "Part b", "definition": "[2,3,4,2,3,4]", "description": "", "templateType": "anything", "can_override": false}, "rt": {"name": "rt", "group": "Part b", "definition": "rtlist[index2]", "description": "", "templateType": "anything", "can_override": false}, "index2": {"name": "index2", "group": "Ungrouped variables", "definition": "random(0..5)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Part c", "definition": "c1list[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Part c", "definition": "c2list[index1]", "description": "", "templateType": "anything", "can_override": false}, "listc2": {"name": "listc2", "group": "Part c", "definition": "random(2..5)^2", "description": "", "templateType": "anything", "can_override": false}, "listc3": {"name": "listc3", "group": "Part c", "definition": "random(2..5)^3", "description": "", "templateType": "anything", "can_override": false}, "listoflistc": {"name": "listoflistc", "group": "Part c", "definition": "[listc3, listc2]", "description": "", "templateType": "anything", "can_override": false}, "c1list": {"name": "c1list", "group": "Part c", "definition": "[2,3]", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part c", "definition": "listoflistc[index1]", "description": "", "templateType": "anything", "can_override": false}, "c2list": {"name": "c2list", "group": "Part c", "definition": "[3,2]", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Part e", "definition": "random(2..9 except d3)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d4": {"name": "d4", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d5": {"name": "d5", "group": "Part e", "definition": "random(2..9 except d4)", "description": "", "templateType": "anything", "can_override": false}, "d6": {"name": "d6", "group": "Part e", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "sold": {"name": "sold", "group": "Part e", "definition": "d1*d2/d3 + d4/d5*d6", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Part d", "definition": "e1list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Part d", "definition": "e2list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Part d", "definition": "e3list[index2]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Part d", "definition": "[3,2,4,3,3,4]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Part d", "definition": "[25,125,8,4,9,27]", "description": "", "templateType": "anything", "can_override": false}, "e3list": {"name": "e3list", "group": "Part d", "definition": "[2,3,2,4,2,2]", "description": "", "templateType": "anything", "can_override": false}, "f1": {"name": "f1", "group": "Ungrouped variables", "definition": "random(2..8)", "description": "", "templateType": "anything", "can_override": false}, "f2": {"name": "f2", "group": "Ungrouped variables", "definition": "random(2..8 except f1)", "description": "", "templateType": "anything", "can_override": false}, "f3": {"name": "f3", "group": "Ungrouped variables", "definition": "random(2..8 except f2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["index1", "index2", "f1", "f2", "f3"], "variable_groups": [{"name": "Part a", "variables": ["a", "esp", "espl", "list2", "list3", "listoflist"]}, {"name": "Part b", "variables": ["b", "blist", "rt", "rtlist"]}, {"name": "Part c", "variables": ["c1", "c2", "listc2", "listc3", "listoflistc", "c1list", "c2list", "c"]}, {"name": "Part e", "variables": ["d1", "d2", "d3", "d4", "d5", "d6", "sold"]}, {"name": "Part d", "variables": ["e1", "e2", "e3", "e1list", "e2list", "e3list"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Utrykk $2^{\\var{f1}/\\var{f2}}\\cdot 2^{\\var{1}/\\var{f3}}$ kan omskrives til $2^{x},$ hvor $x$ er en br&oslash;k. Skriv $x$ i boksen:</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Alle regneregler for potenser gjelder ogs&aring; n&aring;r eksponentene er br&oslash;k!</p>\n<p>Husk regneregelen for potenser:</p>\n<p>$x^a \\cdot x^b = x^{a+b}.$&nbsp;</p>"}], "minValue": "f1/f2 + 1/f3", "maxValue": "f1/f2 + 1/f3", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Utrykk $(3^{\\var{d1}})^{\\var{d2}/\\var{d3}}\\cdot (3^{\\var{d4}/\\var{d5}})^{\\var{d6}}$ kan omskrives til $3^{x},$ hvor $x$ er en br&oslash;k. Skriv $x$ i boksen:</p>", "stepsPenalty": 0, "showStepsLabel": "", "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "adaptiveMarkingUseCondition": "", "adaptiveMarkingNotUsedMessage": "", "exploreObjective": null, "prompt": "<p>Alle regneregler for potenser gjelder ogs&aring; n&aring;r eksponentene er br&oslash;k!</p>\n<p>Husk regneregler for potenser:</p>\n<p>$x^a \\cdot x^b = x^{a+b}.$&nbsp;</p>\n<p>$(x^a)^b=x^{a\\cdot b}$</p>\n<p>Forkort n&aring;r mulig.</p>"}], "minValue": "sold", "maxValue": "sold", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "typeendtoleave": false, "startpassword": "", "autoSubmit": true, "allowAttemptDownload": false, "downloadEncryptionKey": "", "showresultspage": "oncompletion"}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"enterreviewmodeimmediately": true, "showactualmarkwhen": "always", "showtotalmarkwhen": "always", "showanswerstatewhen": "always", "showpartfeedbackmessageswhen": "always", "showexpectedanswerswhen": "inreview", "showadvicewhen": "inreview", "allowrevealanswer": true, "intro": "", "end_message": "", "results_options": {"printquestions": true, "printadvice": true}, "feedbackmessages": [], "reviewshowexpectedanswer": true, "showanswerstate": true, "reviewshowfeedback": true, "showactualmark": true, "showtotalmark": true, "reviewshowscore": true, "reviewshowadvice": true}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "extensions": [], "custom_part_types": [], "resources": []}